<a href="https://colab.research.google.com/github/toche7/MLPython1Day/blob/main/Lab01_ImportantLib_Student.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Important Python Libraries for Data Science


Revised: T. Samanchuen <br>
Version: 001

# Numpy



```python
import numpy as np
```

```python
np.__version__
```

## การสร้าง Array 1 มิติ

``` python
a = np.array([1, 2, 3])
print(a)
```



```python

```

```python

```

## การสร้าง array 2 มิติ


``` python
b = np.array([[1, 2],[3, 4]])
print(b)
```



```python

```

```python

```

## Method ที่อยู่ใน numpy

* np.zeros()
* np.ones()
* np.full()
* np.eye()
* np.random.random()




```python

```

```python

```

## โอเปอร์เรชั่นทางคณิตศาสตร์ของ Array



### Element-wise operations
### การบวก ลบ คูณ หาร  "รายตัว" กรณี Array ทั้งสองมีขนาดเท่ากัน
การบวก ลบ คูณ หารในสมาชิกแต่ละตัวของ array สอง array นั่นคือบวก ลบ คูณ หารเฉพาะสมาชิกในตำแหน่งเดียวกัน สามารถทำได้อย่างตรงไปตรงมาในกรณีที่ array ทั้งสองมีขนาดเท่ากัน

```python
x = np.array([[1,2],[3,4]])
print(x)
```

```python
y = np.array([[5,6],[7,8]])
print(y)
```

```python

```

```python

```

**การยกกำลัง**
มีสองวิธีคือ ใช้เครื่องหมาย  ** หรือ ฟังก์ชัน np.power()

```python

```

```python

```

**ถอดรากที่สอง**
มีสองวิธีคือเรียกฟังก์ชั่น sqrt ของ numpy หรือ ยกกำลังด้วย 1/2

```python

```

```python

```

## ความเร็วของ NumPy

เราสามารถตรวจสอบความเร็วด้วยคำสั่ง %timeit


```python

```

```python

```


## Exercise 1
หากมีข้อมูลนำเข้าเป็น Array ของรัศมีของลูกโป่งที่บรรจุแก๊ส เช่น $r$ = [2, 3, 4, 5] จงเขียนโปรแกรมเพื่อคำนวณหาปริมาณแก๊ส ทั้งหมดของลูกโป่ง
สมมุติว่าลูกโป่งมีลักษณะเป็นทรงกลม โดยปริมาตรของทรงกลมมีสูตรเป็น
$$
V = \frac{4}{3}\pi r^3
$$
Hint: <br>
 $\pi$ ใน Numpy ใช้เป็น np.pi

```python
# Install Nose Test Package for colab
!pip install nose
```

<!-- grade_id: a6a626, points: 0, grade: false, solution: true -->
```python
def Spherical(r):  # modify function name and parameters
    """
    Compute the sum of the Volumn of Spherical.
    r : array of radius
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    return [np.sum(volumn)]
```

<!-- grade_id: e14f8b, points: 1, grade: true, solution: false -->
```python
# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(Spherical(np.array([1])), [4.1887902047863905])
assert_equal(Spherical(np.array([2,3])), [146.60765716752366])
```

## Array indexing
Numpy มีวิธี index array ได้หลายรูปแบบ

### Slicing
เราสามารถใช้ slicing ได้เช่นเดียวกับ Python list และเลือกใช้บนมิติใดของ array ก็ได้

เช่น สร้าง array 2มิติ ที่มีขนาด (3, 4)

```python
a = np.array([[1,2,3,4], [5,6,7,8], [9,10,11,12]])
print(a)
```

```python

```

```python

```

สังเกตว่าการ index แถวที่สองทำได้สองวิธีคือ ใช้ index 1 หรือใช้ slicing 1:2 ซึ่งใน syntax ของ python จะให้ผลลัพธ์ที่ **แตกต่างกัน**

การใช้ index เป็นเพียงตัวเลข เช่น 1 จะลดมิติของ array ลง 1 มิติ (ดูตัวอย่างข้างล่าง)
ในขณะที่การใช้ slicing index เช่น 1:2 เราจะได้ผลลัพธ์เป็น array ที่มีมิติเท่าเดิม

```python
row_r1 = a[1, :]    # ผลลัพธ์คือ array 1 มิติ
row_r2 = a[1:2, :]  # ผลลัพธ์คือ array 2 มิติ
print(row_r1, row_r1.shape)  # พิมพ์ "[5 6 7 8] (4,)"
print(row_r2, row_r2.shape)  #พิมพ์ "[[5 6 7 8]] (1, 4)"

```

## Exercise 2
หากมีข้อมูลที่เป็น array 2 มิติ จงเขียนโปรแกรมที่สามารถตัดเอาเฉพาะข้อมูล 2 Column สุดท้ายของ Array เพื่อนำไปใช้ต่อ

<!-- grade_id: ab1671, points: 0, grade: false, solution: true -->
```python
def slicerLastTwoColumn(a):  # modify function name and parameters
    """
    Compute the squares of the numbers from 1 to n.  [replace with function description]
    """
    s = a.shape
    if s[1] < 3: raise ValueError("n must be at least 3")
    # YOUR CODE HERE
    raise NotImplementedError()

    return b
```

<!-- grade_id: 32ca09, points: 1, grade: true, solution: false -->
```python
a = np.array([[1,2,3,4], [5,6,7,8], [9,10,11,12]])
b = np.array([[1,2,3], [5,6,7], [9,10,11]])

# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
# from nose.tools import assert_array_equal
np.testing.assert_array_equal(slicerLastTwoColumn(a), np.array([[3,  4],[7,  8],[11, 12]]))
np.testing.assert_array_equal(slicerLastTwoColumn(b), np.array([[2,  3],[6,  7],[10, 11]]))
```

## Boolean indexing
เราสามารถ index ค่าใน array ด้วยตัวแปรบูลเลียนได้ ซึ่งถ้าใช้คล่องแล้วจะมีประโยชน์อย่างมาก

```python
a = np.array([[1,2], [3, 4], [5, 6]])

```

```python
bool_idx = (a > 1)   # รีเทิร์นเป็น array ขนาดเดียวกับ a
                     # ค่า bool_idx[i,j] มีค่า เป็น True ถ้า a[i,j] > 2
                     # นอกนั้นจะมีค่าเป็น False

```

```python

```

```python

```

#Pandas

ไลบรารี Pandas มาจากคำว่า Panel Data ถูกสร้างเพื่อกำหนดและวิเคราะห์ข้อมูลในรูปแบบตาราง และถือว่าเป็นไลบรารีหลักที่ใช้ในการทำงานด้าน Data Science
การเรียกใช้ดำเนินการดังนี้

```python
import pandas as ps
```




```python
import pandas as pd
pd.__version__
```

## จัดการข้อมูลแบบอนุกรมด้วย Pandas

การสร้างออบเจ็กข้อมูลแบบอนุกรม

```python
pd.Series(data, index, dtype)
```


```python
data1 = [70,85,60,50,75]
sr1 = pd.Series(data1)
```

```python
sr2 = pd.Series([1.23,111,12.3])
print(sr2)
```

```python
sr3 = pd.Series([1.23,111,12.3,'aa'])
print(sr3)
```

```python

```

```python

```

## การจัดการข้อมูลแบบตารางด้วย Pandas

การสร้างออบเจ็กข้อมูลแบบตาราง

```python
pd.Dataframe(data, index, columns, dtype)
```

```python
data1 = [[1,2,3],
         [2,3,4],
         [3,4,5],
         [4,5,6],
         [5,6,7]]
```

```python

```

```python

```

## Dictionary to Dataframe

```python
data2 = { 'col1' : [1,2,3,4,5],
         'col2' : [2,3,4,5,6],
         'col3' : [3,4,5,6,7],
         'col4' : [4,5,6,7,8]
}
```

```python
df = pd.DataFrame(data2)
```

```python

```

```python

```

## Column และ Index



```
df.columns = ['colName1', 'colName2', 'colName3']
df.index = ['rowName1','rowName2','rowName3']
```



```python
display(df)
```

```python

```

```python

```

# Matplotlib
Matplotlib เป็นไลบรารีสำหรับการพล็อดและจัดการรูปภาพ ซึ่งมีการใช้งานคล้ายกับบน MATLAB.




## Plotting

ตัวอย่างการ plot
``` python
import matplotlib.pyplot as plt
x = np.arange(0, 3 * np.pi, 0.1)
y_sin = np.sin(x)
plt.plot(x, y_sin)
plt.xlabel('x axis label') # กำหนดชื่อบนแกน x
plt.ylabel('y axis label') # กำหนดชื่อบนแกน y
plt.title('Sine and Cosine') # กำหนดชื่อไตเติ้ล
plt.legend(['Sine', 'Cosine']) # กำหนดลาเบลของกราฟแต่ละเส้น (ตามลำดับคำสั่ง plt.plot)
plt.show() # เมื่อระบุสิ่งที่เราต้องการพล็อตครบแล้ว สามารถแสดงผลการพล็อตได้ด้วยคำสั่ง plt.show()
```



```python

```

```python

```

## Subplot
เราสามารถแบ่งรูปภาพเป็นหลายส่วน และพล็อตแต่ละกราฟในแต่ละส่วนย่อยได้ด้วยฟังก์ชัน `supplot`




``` python
import matplotlib.pyplot as plt
x = np.arange(0, 3 * np.pi, 0.1)
y_sin = np.sin(x)
fig, ax = plt.subplots()
ax.plot(x, y_sin)
ax.set_title('Title')
ax.set_xlabel('X')
ax.set_ylabel('Y')
```



```python

```

ตัวอย่างการใช้ subplot

``` python
x = np.arange(0, 3 * np.pi, 0.1)
y_sin = np.sin(x)
y_cos = np.cos(x)
fig, (ax1,ax2) = plt.subplots(2,1, constrained_layout = True)
ax1.plot(x, y_sin),
ax1.set_title('Sin')
ax1.set_xlabel('X')
ax1.set_ylabel('Y')

ax2.plot(x, y_cos, 'r')
ax2.set_title('Cos')
ax2.set_xlabel('X')
ax2.set_ylabel('Y')
```



```python

```

```python

```

```python
import matplotlib.pyplot as plt
import numpy as np

# Create a Figure with 2 rows and 2 columns of subplots:
fig, ax = plt.subplots(2, 2, constrained_layout = True)
#fig.tight_layout(pad=2.0)

x = np.linspace(0, 5, 100)

# Index 4 Axes arrays in 4 subplots within 1 Figure:

ax[0, 0].plot(x, np.sin(x), 'g') #row=0, column=0
ax[0, 0].set_xlabel('x1')
ax[0, 0].set_ylabel('y1')
ax[0, 0].set_title('Ax[0,0]')

ax[1, 0].plot(range(100), 'b') #row=1, column=0
ax[1, 0].set_xlabel('x1')
ax[1, 0].set_ylabel('y1')
ax[1, 0].set_title('Ax[1,0]')

ax[0, 1].plot(x, np.cos(x), 'r') #row=0, column=1
ax[0, 1].set_xlabel('x1')
ax[0, 1].set_ylabel('y1')
ax[0, 1].set_title('Ax[0,1]')

ax[1, 1].plot(x, np.tan(x), 'k') #row=1, column=1
ax[1, 1].set_xlabel('x1')
ax[1, 1].set_ylabel('y1')
ax[1, 1].set_title('Ax[1,1]')



plt.show()
```

# Seabon

```python
# Import seaborn
import seaborn as sns

# Apply the default theme
sns.set_theme()

# Load an example dataset
tips = sns.load_dataset("tips")

# Create a visualization
sns.relplot(
    data=tips,
    x="total_bill", y="tip", col="time",
    hue="smoker", style="smoker", size="size",
)
```

```python
penguins = sns.load_dataset("penguins")
sns.pairplot(data=penguins, hue="species")
```

```python
penguins
```